# Teste 2  - Ajuste de parâmetros

Responda as questões abaixo.

Não crie novas células nem apague células existentes. Apenas responda nas células indicadas.

Escreva o texto em Markdown e LaTeX, conforme instruções em [Jupyter - Markdown cells](https://jupyter-notebook.readthedocs.io/en/stable/examples/Notebook/Working%20With%20Markdown%20Cells.html), e escreva os códigos em [Julia](https://julialang.org).

Apenas os pacotes indicados na célula abaixo podem ser utilizados para responder às questões cujas respostas requerem códigos.

```julia
using LinearAlgebra
using Statistics
using Optim
using Plots
```

Output:
```
┌ Info: Precompiling Plots [91a5bcdd-55d7-5caf-9e0b-520d859cae80]
└ @ Base loading.jl:1423
```

## Questão 1

Suponha que tenhamos uma amostra de três medições $(x_i, y_i)$, $i=1,2,3,4$, e que queiramos determinar três parâmetros $\beta = (\beta_0, \beta_1, \beta_2)$ para que o modelo
$$ \hat y = \beta_0 + \beta_1 x + \beta_2 x^2
$$
melhor se ajuste aos dados no sentido de minimizar a soma do quadrados dos resíduos $r_i = y_i - \hat y_i$, onde $\hat y_i = \beta_0 + \beta_1 x_i + \beta_2 x_i^2$.

Escreva, abaixo, uma condição nos dados da amostra que garanta que exista uma e somente uma solução para esse problema, ou seja, que exista um, e somente um, conjunto de parâmetros que minimize a soma dos quadrados dos resíduos.

<!-- grade_id: cell-6116687273053247, points: 4, grade: true, solution: true -->
Para que um sistema linear $A\mathbf{x} = \mathbf{b}$ qualquer tenha solução única, é necessário que seu núcleo seja trivial (garante a unicidade, mas não a existência), e que, denotando por $A_{*j}$ os vetores coluna de $A$, $b \in \mathrm{span}(A_{*j})$ -- o que garante a existência da solução. À primeira condição equivale dizer que $\det A \neq 0$; $A$ e $A^T$ são invertíveis; os $A_{*j}$ são linearmente independentes. Considerando que são três os parâmetros a se determinar, a matriz de Vandermonde associada a esse problema de mínimos quadrados tem que ter três colunas. Do contrário, o problema não está bem definido, o que se evidencia se considerada a forma normal do mesmo:
\begin{equation}
A^TA\boldsymbol{\beta} = A^T\mathbf{y}
\end{equation}
Em particular, essas colunas têm de ser linearmente independentes, para que esteja garantida a unicidade da solução. Observe que, também devido a $(1)$, $A$ tem de ter três linhas, porquanto são três as amostras. Logo $A$ é quadrada, e, portanto, seu determinante é dado pela fórmula:
\begin{equation}
\det A = \prod_{i<j} (x_j - x_i)
\end{equation} 
que é diferente de zero se, e somente se, $x_i \neq x_j$, para $i, j \in \{1, 2, 3\}$, o que garante a independencia linear das colunas, a invertibilidade de $A$ e, por conseguinte, a unicidade da solução.

## Questão 2

Considere o modelo da forma
$$ y = \frac{\alpha x}{(1 + \beta x)^3}.
$$

Ache uma mudança de variáveis que transforme esse modelo em um modelo linear nos (novos) parâmetros. Justifique o seu desenvolvimento, exibindo a mudança de variáveis e o modelo linear.

<!-- grade_id: cell-c41d0d4572e6b45a, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 3

Em um ensaio de laboratório, foram testadas diferentes doses de uma nova droga em um certos indivíduos voluntários. Para cada nível de dosagem, 250 indivíduos foram expostos à droga e 250, a um placebo. O número de sucessos foi reportado:

| dosagem | sucesso droga | sucesso placebo
|:---:|:---:|:---:|
| 1 | 29 | 5 |
| 2 | 56 | 7 |
| 3 | 85 | 8 |
| 4 | 112 | 12 |
| 5 | 134 | 11 |
| 6 | 143 | 17 |

Use o método de mínimos quadrados para encontrar os parâmetros $\beta = (\beta_0, \beta_1)$ que melhor ajustam o modelo linear $y = \beta_0 + \beta_1 x$ aos dados acima de sucesso da droga, no sentido de minimizar a soma dos quadrados dos resíduos, onde $x=1,\ldots, 6$ indica a dosagem e $y$ indica o número de sucesso nos que receberam a droga.

1. Escreva, abaixo, o código para encontrar o vetor bidimensional $\beta = [\beta_0, \beta_1]$ com esses parâmetros;
2. Complete a definição da função $f(x)$ que retorna a estimativa do número de sucesso, para cada dosagem $x$, onde $x$ deve ser um número em ponto flutuante `Float64`.
3. Use o pacote `Plots` para exibir um gráfico dos dados e do ajuste por esse modelo linear, com os parâmetros encontrados.

A resposta de cada item acima deve estar em cada uma das células de código abaixo, conforme indicado.

<!-- grade_id: cell-2e6ed5733a5f6de9, points: 0, grade: false, solution: true -->
```julia
# Resposta para o item 1 da questão 4 aqui
yi = [] # número de sucessos
β = [] # parâmetros
# YOUR CODE HERE
println("Os parâmetros do melhor ajuste são β₀ = $(round(β[1], digits=3)) e β₁ = $(round(β[2], digits=3))")
```

<!-- grade_id: cell-5bbb0cff3099cfdf, points: 2, grade: true, solution: false -->
```julia
@assert yi isa Vector
@assert β isa Vector
@assert length(β) == 2
```

<!-- grade_id: cell-8eb77fac9ed32b43, points: 0, grade: false, solution: true -->
```julia
# Resposta para o item 2 da questão 4 aqui
function f(x, β)
    y = nothing
# YOUR CODE HERE
    return y
end
```

<!-- grade_id: cell-80e46c56c03f59ad, points: 1, grade: true, solution: false -->
```julia
@assert f(1.0, [1.0, 2.0]) isa Float64
```

<!-- grade_id: cell-eb5b11701235f8c4, points: 1, grade: true, solution: true -->
```julia
# Resposta para o item 3 da questão 4 aqui
# Use `plot!` e `scatter!` para traçar o ajuste e os dados "em cima" da figura criada a seguir
# Use o argumento `label` para indicar os traçados de ajuste e de dados
plot(ylim = (0.0, 250), title="Dados e modelo linear ajustado", titlefont=10, legend=:topleft)
# YOUR CODE HERE
```

## Questão 4

Queremos modelar um problema com valores escalares $z$ dependentes de duas variáveis $x$ e $y$ através de uma função da forma

$$ z = \frac{ \alpha xy }{ 1 + \gamma x + \delta y }
$$

Foram feitas as seguintes medições.

| x | y | z |
| :--: | :---: | :---: |
| 9.1 | 7.0 | 12.75 |
| 0.7 | 1.2 | 1.11 |
| 6.3 | 1.7 | 2.60 |
| 5.7 | 3.2 | 4.59 |
| 3.7 | 6.6 | 6.98 |
| 2.7 | 3.7 | 3.69 |
| 7.3 | 0.1 | 0.23 |
| 2.3 | 7.5 | 4.90 |
| 5.8 | 0.2 | 0.46 |
| 1.0 | 4.3 | 2.51 |
| 1.3 | 6.2 | 2.95 |
| 3.6 | 1.5 | 2.64 |
| 4.2 | 0.4 | 0.75 |
| 5.8 | 7.4 | 7.47 |
| 4.7 | 5.3 | 7.83 |
| 5.7 | 3.7 | 4.73 |
| 1.4 | 0.8 | 1.42 |
| 1.7 | 9.5 | 4.90 |
| 9.4 | 1.8 | 2.50 |
| 0.4 | 3.6 | 0.79 |

Queremos encontrar parâmetros $\alpha$, $\gamma$ e $\delta$ que melhor ajustam o modelo aos dados indicados, no sentido de minimizar a soma dos quadrados dos resíduos.
 
Assim como no modelo de Michaelis-Menten, a função acima pode ser transformada em um modelo linear em termos de parâmetros apropriados relacionados com os parâmetros acima e com mudanças adequadas nas variávels $x$, $y$ e $z$. Isso nos permite fazer o ajuste usando o método de mínimos quadrados. Também podemos fazer o ajuste do modelo na forma não linear, usando `Optim.jl`, por exemplo.

Abaixo, temos uma célula de Markdown e duas células de código.

1. Na primeira, deduza o modelo linear nos parâmetros conforme mencionado acima, indicando a relação das novas variáveis e dos novos parâmetros com as variáveis e parâmetros originais.

1. Na segunda, obtenha os parâmetros $\alpha=\alpha_q$, $\delta=\delta_q$ e $\gamma=\gamma_q$  de melhor ajuste via mínimos quadrados usando a transformação mencionada acima e calcule a soma dos quadrados dos resíduos `rss_q` para o modelos ajustado.

1. Na terceira, trabalhe diretamente com o modelo não-linear, obtenha os parâmetros $\alpha=\alpha_\textrm{opt}$, $\delta=\delta_\textrm{opt}$ e $\gamma=\gamma_\textrm{opt}$ de melhor ajuste via `Optim.jl` e calcule a soma dos quadrados dos resíduos `rss_opt` para o modelos ajustado.

<!-- grade_id: cell-f61232bac77569ea, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-2d2aa9f42ffc2b84, points: 0, grade: false, solution: true -->
```julia
# Código para a sua resposta para o item 2 da questão 5, resolvendo via minimos quadrados, aqui nesta célula

xyz = [
    9.1 7.0 12.75;
    0.7 1.2 1.11;
    6.3 1.7 2.60;
    5.7 3.2 4.59;
    3.7 6.6 6.98;
    2.7 3.7 3.69;
    7.3 0.1 0.23;
    2.3 7.5 4.90;
    5.8 0.2 0.46;
    1.0 4.3 2.51;
    1.3 6.2 2.95;
    3.6 1.5 2.64;
    4.2 0.4 0.75;
    5.8 7.4 7.47;
    4.7 5.3 7.83;
    5.7 3.7 4.73;
    1.4 0.8 1.42;
    1.7 9.5 4.90;
    9.4 1.8 2.50;
    0.4 3.6 0.79
]

α_q, γ_q, δ_q = nothing, nothing, nothing

rss_q = nothing

# YOUR CODE HERE

println("Solução via mínimos quadrados:")
println("    Parâmetros:
        α_q = $(round(α_q, digits=3))
        γ_q = $(round(γ_q, digits=3))
        δ_q = $(round(δ_q, digits=3))")
println("    Soma dos quadrados dos resíduos:
        rss_q = $(round(rss_q, digits=3))")
```

<!-- grade_id: cell-cc7be0f3d5030892, points: 2, grade: true, solution: false -->
```julia
@assert isassigned(Ref(rss_q))

```

<!-- grade_id: cell-c2e8ac6283da3106, points: 0, grade: false, solution: true -->
```julia
# Código para a sua resposta para o item 3 da questão 5, resolvendo via Optim, aqui nesta célula

α_opt, γ_opt, δ_opt = nothing, nothing, nothing

rss_opt = nothing

# YOUR CODE HERE

println("Solução via \"Optim.jl\":")
println("    Parâmetros:
        α_opt = $(round(α_opt, digits=3))
        γ_opt = $(round(γ_opt, digits=3))
        δ_opt = $(round(δ_opt, digits=3))")
println("    Soma dos quadrados dos resíduos:
        rss_opt = $(round(rss_opt, digits=3))")
```

<!-- grade_id: cell-56e293d076f305d2, points: 2, grade: true, solution: false -->
```julia
@assert isassigned(Ref(rss_opt))

```